## Assignment II
In this assignment, we are going practice the pandas syntax.

The structure of the assignment is written in terms of functions. This mean that you can validate your assignment at any time. You will contruct each answer in the form of a function. Each of these questions already has the prebuilt function syntax with doc string included that tells you what the function must accept and return. 

```python
import pandas as pd
df = pd.read_csv('../../data/ChiCrimes2001-present.csv')
```

### Q1 -- Shape
What is the shape of this data? Using the cell below, fill in the code that will make the function `dfShape` take a dataframe as input and output the shape of that dataframe.

<!-- grade_id: shape, points: 0, grade: false, solution: true -->
```python
def dfShape(aDataframe):
    """Accepts any dataframe.
    
    Parameters
    ----------
    aDataframe : variable
        The name of the dataframe
        
    Returns
    -------
    The shape of the dataframe. 
    
    """
    # YOUR CODE HERE
    
    raise NotImplementedError()
```

<!-- grade_id: test_shape, points: 1, grade: true, solution: false -->
```python
assert dfShape(df) == (100000, 22)
```

### Q2 -- Find a Row
In this question, the task is to subset a particular row. Fill in your function below. It should, for example, be able to accept a dataframe and return row #8. 

**Hint**: perhaps look into the `.tolist()` method of a pandas column

<!-- grade_id: row, points: 0, grade: false, solution: true -->
```python
def dfRownum(aDataframe, rownum):
    """Accepts a dataframe and row.
    
    Parameters
    ----------
    aDataframe : variable
        The name of the dataframe
    rownum : int
        The row number to return.
        
    Returns
    -------
    A list containing the row specified.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_row, points: 2, grade: true, solution: false -->
```python
assert dfRownum(df, 13) == [6122609, 'HP216895', '03/09/2008 10:00:00 PM', '042XX N ALBANY AVE', '0610', 'BURGLARY', 'FORCIBLE ENTRY', 'RESIDENCE-GARAGE', False, False, 1724, 17.0, 33.0, 16.0, '05', 1154956.0, 1928079.0, 2008, '02/04/2016 06:33:39 AM', 41.958463250999998, -87.705714737000008, '(41.958463251, -87.705714737)']
assert dfRownum(df, 72447) == [10484721, 'HZ224442', '04/13/2016 01:20:00 PM', '033XX W 63RD PL', '1020', 'ARSON', 'BY FIRE', 'STREET', False, False, 823, 8.0, 15.0, 66.0, '09', 1155321.0, 1862317.0, 2016, '04/20/2016 03:55:16 PM', 41.777997988000003, -87.706139267999987, '(41.777997988, -87.706139268)']
```

### Q3 -- Number of Crimes
How many crimes of a particular crime type are there in this dataset? To answer this question, create a function that creates a pandas column of crime occurances.

**Hint**: think about what `.value_counts()` method does. 

<!-- grade_id: PrimeCrimes, points: 0, grade: false, solution: true -->
```python
def dfPrmCrimes(aDataframe, crime):
    """Reads in the dataset file and returns an integer number of occurances for particular crime from variable primary type.
    
    Parameters
    ----------
    aDataframe : variable
        The name of the dataframe
    crime : string
        The crime selected from the column "Primary Type"
        
    Returns
    -------
    An integer for the number of occurances of this crime.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_PrmCrimes, points: 5, grade: true, solution: false -->
```python
assert dfPrmCrimes(df, 'THEFT') == 20734
assert dfPrmCrimes(df, 'HOMICIDE') == 116
assert dfPrmCrimes(df, 'NARCOTICS') == 11234
assert dfPrmCrimes(df, 'ROBBERY') == 3737
assert dfPrmCrimes(df, 'LIQUOR LAW VIOLATION') == 231
```

### Q4 -- Crime by Location
In this question, we want to figure out how many occurances of a particular crime type there are at a particular type of location. For example, how many occurances of 'THEFT' were there in an alley? To answer this question, fill in the code in the function below that will take as input a dataframe, a crime and a location. 

<!-- grade_id: generalCrime, points: 0, grade: false, solution: true -->
```python
def dfCrimeLoc(aDataframe, crime, location):
    """Reads in the dataset file and returns an integer number of occurances for particular crime from variable primary type.
    
    Parameters
    ----------
    aDataframe : variable
        The name of the dataframe
    crime : string
        The crime selected from the column "Primary Type"
    location : string
        The location type selected from the "Location Description" column (e.g. 'STREET')
        
    Returns
    -------
    An integer for the number of occurances of this crime at this location type.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_crimeloc, points: 5, grade: true, solution: false -->
```python
assert dfCrimeLoc(df, 'theft', 'street') == 5668
assert dfCrimeLoc(df, 'WEAPONS VIOLATION', 'SCHOOL, PUBLIC, BUILDING') == 90
assert dfCrimeLoc(df, 'GAMBLING', 'SCHOOL, PUBLIC, BUILDING') == 3
assert dfCrimeLoc(df, 'MOTOR VEHICLE THEFT', 'alley') == 90
assert dfCrimeLoc(df, 'KIDNAPPING', 'SIDEWALK') == 15
```

### Q5 -- Crime by Year
Now we want to know how many occurance of a particular crime there were in a particular year. Fill in the code below to create a function that accepts a dataframe, crime type and year and outputs the integer value for number of occurances of that crime in that year. 

<!-- grade_id: crimebyyear, points: 0, grade: false, solution: true -->
```python
def dfCrimeByYear(aDataframe, crime, year):
    """Reads in the dataset file and returns an integer number of occurances for particular crime from variable primary type.
    
    Parameters
    ----------
    aDataframe : variable
        The name of the dataframe
    crime : string
        The crime selected from the column "Primary Type"
    year : year of crime as an integer
    Returns
    -------
    An integer for the number of occurances of this crime at this location type.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_crimebyyear, points: 4, grade: true, solution: false -->
```python
assert dfCrimeByYear(df, 'arson', 2004) == 8
assert dfCrimeByYear(df, 'gambling', 2015) == 3
assert dfCrimeByYear(df, 'assault', 2015) == 241
assert dfCrimeByYear(df, 'robbery', 2001) == 298
```

### Q6 -- Crime over Time
Now we want to modify the question above slightly to find out how many occurances of a particular crime there are over a range of years. Fill in the function below to create a function that accepts as input the dataframe, crime and range of years (in the form of a list of years) and outputs a list of the number of occurances of those crimes for those years. For example, dfCrimerange(df, 'robbery', [2001,2002]) should return [298, 303], meaning that in 2001, there were 298 robberies and in 2002, there were 303.

**Hint**: use a for loop

<!-- grade_id: crime_range, points: 0, grade: false, solution: true -->
```python
def dfCrimerange(aDataframe, crime, range_years):
    """Accepts a dataframe, a crime and a range of years (in the form of a list)

    Parameters
    ----------
    aDataframe : variable
        The name of the dataframe
    crime : string
        The crime selected from the column "Primary Type"
    range_years : range of years
    
    Returns
    -------
    A list of the occurances for a particular crime for the range of years entered.
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: test_crime_range, points: 2, grade: true, solution: false -->
```python
assert dfCrimerange(df, 'theft', range(2002,2008)) == [1593, 1624, 1612, 1344, 1443, 1418]
assert dfCrimerange(df, 'robbery', [2001,2002]) == [298, 303]
```